---
# Welcome to CE880
### This is your week-4 : Problem notebook

For this problem set, we'll be using the Jupyter notebook and please upload this notebook to [Google Colab](https://colab.research.google.com/). 

### Task: Complete all the three questions given as follows:

```python
# install sklearn package if you are using local computer and if you are using Google Coalb then skip this step
!pip install -U scikit-learn
```

```python
import pandas as pd
import numpy as np
from sklearn.cluster import DBSCAN
from sklearn.cluster import KMeans
from sklearn.datasets import make_circles, make_moons
from sklearn.metrics import adjusted_rand_score
from sklearn.preprocessing import StandardScaler
from sklearn.metrics import silhouette_score, adjusted_rand_score
from sklearn.decomposition import PCA
scaler = StandardScaler()
```

---
## Question 1: 
We have already loaded `moon` data for you using `make_moons` function. 

Write a code to perform `k-means` clustering on `scaled_features` variable and return `silhouette_score` as an output

<!-- grade_id: myScore_kmean_ans, points: 0, grade: false, solution: true -->
```python
# Let's make two moons, where n is the number of samples
features, true_labels = make_moons(n_samples=250, noise=0.05, random_state=42)
# fit_transform() method is basically the combination of fit method and transform method, 
# it is equivalent to fit().transform(). This method performs fit and transform on the 
# input data at a single time and converts the data points. 
# If we use fit and transform separate when we need both then it 
# will decrease the efficiency of the model so we use fit_transform() 
# which will do both the work.
scaled_features = scaler.fit_transform(features)
myscore = []

def myScore_kmean(scaled_features):
    """Write code to perform k-means clustering and print silhouette_score"""
    # YOUR CODE HERE
    km = KMeans(n_clusters=6, random_state=42)
    km.fit_predict(scaled_features)
    score = silhouette_score(scaled_features, km.labels_, metric='euclidean')#calculate the silhouette_score
    return float('%.3f' %score)#Return Score value using float
    raise NotImplementedError()
```

<!-- grade_id: myScore_kmean_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
import math
assert math.isclose(myScore_kmean(scaled_features), 0.5, rel_tol=0.10)
```

---
## Question 2: 

Write a code to perform `DBSCAN` clustering algorithm on `scaled_features` and return `silhouette_score` as an output. Please set the `eps` = 0.2 and `min_samples` = 5

<!-- grade_id: myScore_DBSCAN_ans, points: 0, grade: false, solution: true -->
```python

# Let's make two moons, where n is the number of samples
from sklearn.datasets import make_moons
import numpy as np
from sklearn.cluster import DBSCAN
from sklearn import metrics
import math
features, true_labels =  make_moons(n_samples=250, noise=.1)

def myScore_DBSCAN(scaled_features):
    """Write code to perform DBSCAN algorithm and print silhouette_score"""
    # YOUR CODE HERE
    db = DBSCAN(eps=0.2, min_samples=5).fit(scaled_features)
    labels = db.labels_
    score = silhouette_score(scaled_features, labels, metric='euclidean')#calculate the silhouette_score
    return float('%.3f' %score)
    raise NotImplementedError()
```

<!-- grade_id: myScore_DBSCAN_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
assert math.isclose(myScore_DBSCAN(scaled_features), 0.2, rel_tol=0.10)
```

## Question 3: 

Complete this code perform `PCA`  on `iris data` and return `explained_variance_ratio_` as an output for the first two principal components (PCs)

<!-- grade_id: myPCA_ans, points: 0, grade: false, solution: true -->
```python
import pandas as pd
import numpy as np
import pylab as pl
import math
from sklearn import datasets
from sklearn.model_selection import train_test_split
from sklearn.decomposition import PCA
url = "https://archive.ics.uci.edu/ml/machine-learning-databases/iris/iris.data"
# load dataset into Pandas DataFrame
df = pd.read_csv(url, names=['sepal length','sepal width','petal length','petal width','target'])

features = ['sepal length', 'sepal width', 'petal length', 'petal width']
# Separating out the features
x = df.loc[:, features].values
# Separating out the target
y = df.loc[:,'target'].values


# Standardizing the features
x = StandardScaler().fit_transform(x)

x_train, x_test, y_train, y_test = train_test_split(x, y, test_size=0.2, stratify=y, random_state=0)
from sklearn.preprocessing import StandardScaler
sc = StandardScaler()
x_train_std = sc.fit_transform(x_train)
x_test_std = sc.transform(x_test)
  
def myPCA(x):
    """Complete this code return the explained variance ratio for the first two prinical 
    components after doing PCA on x """
    # YOUR CODE HERE
    pca = PCA(n_components=2)
    x_train_pca = pca.fit_transform(x)
    #return pca explained variance ratio in list
    return list(pca.explained_variance_ratio_)
    raise NotImplementedError()
```

<!-- grade_id: myPCA_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
x= myPCA(x)
y = [0.7277045209380135, 0.23030523267680633]
np.testing.assert_allclose(x, y, rtol=1e-5, atol=0)

```